Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Zhongzhi Zhang (Nino)"
COLLABORATORS = ""
```

---

# Newton's method for rational functions

Newton's method solves $f(x) = 0$ using a first order Taylor approximation at each point.  We can use it to compute square roots and related functions.  Sample code for Newton's method is provided.

```python
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('ggplot')
import numpy as np

def newton(func, x, verbose=False):
    """Solve f(x) = 0 using initial guess x.
    
    The provided function func must return a pair of values,
    f(x) and its derivative f'(x).  For example, to solve
    the equation x^2 - 3 starting from initial guess x=1,
    one would write
    
    def func(x):
        return x**2 - 3, 2*x
        
    newton(f, 1)
    """
    for i in range(100):
        fx, dfx = func(x)
        if verbose:
            print(func.__name__, i, x, fx)
        if np.abs(fx) < 1e-12:
            return x, fx, i
        try:
            x -= fx / dfx
        except ZeroDivisionError:
            return x, np.NaN, i
```

## Square root

We can estimate $\sqrt{a}$ as the root of $f(x) = x^2 - a$.

```python
def my_sqrt(a, guess=1, verbose=False):
    """Compute sqrt(a) using Newton's method.
    """
    def func(x):
        return x**2 - a, 2*x
    value, residual, numits = newton(func, guess, verbose)
    return value

my_sqrt(4, guess=1, verbose=True)
```

Output:
```
func 0 1 -3
func 1 2.5 2.25
func 2 2.05 0.20249999999999968
func 3 2.000609756097561 0.002439396192741583
func 4 2.0000000929222947 3.716891878724482e-07
func 5 2.000000000000002 8.881784197001252e-15
2.000000000000002
```

* Experiment to see if there are positive values of `a` and/or `guess` for which this does not converge.

## Reciprocal square root

Similar to the above, express the reciprocal square root $1/\sqrt{a}$ as a rootfinding problem.

<!-- grade_id: cell-d432be9e53e38a15, points: 0, grade: false, solution: true -->
```python
def my_rsqrt(a, guess=.5, verbose=False):
    """Compute 1/sqrt(a) using Newton's method.
    """
    def func(x):
        # YOUR CODE HERE
        return 1/x**2 - a, -2/x**3
        raise NotImplementedError()
    value, residual, numits = newton(func, guess, verbose)
    return value

my_rsqrt(9, guess=.5, verbose=True)
```

Output:
```
func 0 0.5 -5.0
func 1 0.1875 19.444444444444443
func 2 0.2515869140625 6.798792811486441
func 3 0.30572039512026095 1.6991967907584549
func 4 0.3299969424357781 0.182906620291984
func 5 0.3332834086900596 0.0026965365279654208
func 6 0.3333333221177783 6.056400039256005e-07
func 7 0.33333333333333276 3.019806626980426e-14
0.33333333333333276
```

<!-- grade_id: cell-91a6158aaab0493e, points: 1, grade: true, solution: false -->
```python
assert np.isclose(1/my_rsqrt(1.05)**2, 1.05)
assert np.isclose(1/my_rsqrt(100, guess=.01)**2, 100)
print("Tests pass")
```

Output:
```
Tests pass
```

### Uniqueness

Find a positive value of `my_guess` for which `my_rsqrt(10, my_guess)` converges to the negative root.

<!-- grade_id: cell-a820b936819b69a8, points: 0, grade: false, solution: true -->
```python
# my_guess = SOME_NUMBER
my_guess = 0.3

# YOUR CODE HERE
#raise NotImplementedError()

my_rsqrt(10, my_guess, verbose=True)
```

Output:
```
func 0 0.3 1.1111111111111107
func 1 0.315 0.07810531620055272
func 2 0.316220625 0.0004516528590485791
func 3 0.31622776577495343 1.5298118327677912e-08
func 4 0.31622776601683794 0.0
0.31622776601683794
```

```python
assert np.isclose(my_rsqrt(10, my_guess), 1/np.sqrt(10))
print("Test pass")
```

Output:
```
Test pass
```

<!-- grade_id: cell-13a86ac1ff4c60b7, points: 1, grade: true, solution: false -->
```python
assert np.isclose(my_rsqrt(10, my_guess), -1/np.sqrt(10))
print("Tests pass")
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-8-64e8d857e4ce> in <module>
----> 1 assert np.isclose(my_rsqrt(10, my_guess), -1/np.sqrt(10))
      2 print("Tests pass")

AssertionError:
```

## Avoiding division (optional)

This implementation performs division in defining the function $f(x)$ and the Newton step,
$$ x_{\text{next}} = x - f(x) / f'(x) . $$
(Division is a relatively expensive operation compared to addition and multiplication.)
Substitute the expressions you used for $f(x)$ and $f'(x)$ and simplify so that there is no division in the Newton step itself.
Use this in a function that always does exactly five iterations and does not test for convergence, etc.

<!-- grade_id: cell-6535d49fe71fd2f4, points: 0, grade: false, solution: true -->
```python
def my_rsqrt_fast(a, guess):
    x = guess
    for i in range(100):
        # Single line updating x to the next iteration,
        # x = ...
        # YOUR CODE HERE
        if np.abs(1/x**2 - a) < 1e-12:
            return x
        x -= -0.5*x
        raise NotImplementedError()
    return x

my_rsqrt_fast(5, .5)
```

<!-- grade_id: cell-058d031451741601, points: 1, grade: true, solution: false -->
```python
assert np.isclose(my_rsqrt(5, .5), 1/np.sqrt(5))
print('Tests pass')
```

```python

```